# Moments of a Distribution
taken from http://www.aip.de/groups/soe/local/numres/bookcpdf/c14-1.pdf


```python
import pandas as pd
import matplotlib.pyplot as plt
import numpy as np
import scipy
import scipy.stats
```

## historical daily prices for S&P 500 TR (total return) index
dataset source https://finance.yahoo.com/quote/%5ESP500TR/history?period1=568278000&period2=1548226800&interval=1d&filter=history&frequency=1d
 
NOTE: Adj Close factors in dividends, so that's a more accurate measure of returns

```python
# dataset source https://finance.yahoo.com/quote/%5ESP500TR/history?period1=568278000&period2=1548226800&interval=1d&filter=history&frequency=1d
# historical daily prices for S&P 500 TR (total return) index
# NOTE: Adj Close factors in dividends, so that's a more accurate measure of returns
url="https://raw.githubusercontent.com/marilynwaldman/kurtosis/master/SP500TR.csv"
df = pd.read_csv(url)
df.head()
```

```python
returns = np.ediff1d(df['Adj Close'].values)  # take Closing value (today) - Closing value (yesterday)
pct_returns = returns / df['Adj Close'].values[:-1]  # convert to % by dividing by Closing value (yesterday)
pct_returns
```

```python
# plot histogram compared to Gaussian.  Looks like this is not a Gaussian!
# in fact it is probably a Laplace distribution
# https://sixfigureinvesting.com/2016/03/modeling-stock-market-returns-with-laplace-distribution-instead-of-normal/
plt.figure(1)
plt.hist(pct_returns, bins=100, density=True)
plt.xlim((min(pct_returns), max(pct_returns)))

# calculate fitted Gaussian
mean = np.mean(pct_returns)
variance = np.var(pct_returns)
sigma = np.sqrt(variance)
x = np.linspace(min(pct_returns), max(pct_returns), 100)
plt.plot(x, scipy.stats.norm.pdf(x, mean, sigma))

plt.show()

```

```python
# returns have LONG tails.  "Outlier" evens definitely happen (i.e. catastrophic market crashes)
# kurtosis=3 would be Gaussian.  Anything larger has fat tails
from scipy.stats import kurtosis

kurtosis(pct_returns)
```

```python
# plot again on log scale to really show those fat tails

plt.figure(1)
plt.hist(pct_returns, bins=20, density=True)
plt.xlim((min(pct_returns), max(pct_returns)))

# calculate fitted Gaussian
mean = np.mean(pct_returns)
variance = np.var(pct_returns)
sigma = np.sqrt(variance)
x = np.linspace(min(pct_returns), max(pct_returns), 100)
plt.plot(x, scipy.stats.norm.pdf(x, mean, sigma))

plt.yscale('log')
plt.show()
```

# Compute the minimum of pct_returns

<!-- grade_id: cell-de82703a1c5d0ce9, points: 0, grade: false, solution: true -->
```python
# pct_rtns_min = np.min(pct_returns)
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-1c506a4ab8c48d76, points: 10, grade: true, solution: false -->
```python
assert pct_rtns_min == np.min(pct_returns)

```

# Compute the maximum of pct_returns

<!-- grade_id: cell-3717299e652a94f0, points: 0, grade: false, solution: true -->
```python
# pct_rtns_max = np.max(pct_returns)
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-2b80b629210d048a, points: 10, grade: true, solution: false -->
```python
assert pct_rtns_min == np.min(pct_returns)

```